---
title: "R para análise de dados em Psicologia: Uma Introdução - Exercícios"
output: 
  learnr::tutorial:
    allow_skip: TRUE
    language: pt
runtime: shiny_prerendered
editor_options: 
  chunk_output_type: console
---

```{r setup, include=FALSE}
library(learnr)
library(tidyverse)
library(gradethis)
gradethis_setup()
knitr::opts_chunk$set(echo = F)

tutorial_options(
  exercise.timelimit = 30,
)
```

## Bem-vindos

![](images/logo_pos-02.png){width="184" height="56"}

![](images/lappsie-02.png){width="185" height="40"}

Este é um conjunto de exercícios que acompanham a oficina "R para análise de dados em Psicologia: Uma Introdução".

A oficina é realizada pelo Laboratório de Práticas e Pesquisas em Psicologia e Educação (LAPPSIE), vinculado ao Programa de Pós-graduação Profissional em Psicologia e Políticas Públicas da Universidade Federal do Ceará.

Siga com as atividades no seu ritmo, mas é interessante que, a cada módulo, você tente responder o bloco de exercícios correspondente.

Também aproveite a oportunidade para discutir com os seus colegas sobre como foram criadas as soluções para as atividades.

Esta página está em construção e passando por inúmeras atualizações.

## Básico do R

### Introdução

Neste bloco de exercícios você vai praticar:

-   Operações matemáticas
-   Criação de vetores com o operador de atribuição `<-`
-   Realizar operações com vetores
-   Utilizar funções básicas do R

### Exercício 1

Encontre o log natural, log na base 10, log na base 2 e a raiz quadrada de `12.43`. Some os resultados os resultados encontrados e arredonde a soma para duas casas decimais.

```{r log-raiz, exercise=TRUE}

```

```{r log-raiz-hint}
log()

log10()

log2()

sqrt()

round()
```

```{r log-raiz-quiz}
quiz(
  question("Qual o valor da soma dos resultados encontrados",
    answer("10.87"),
    answer("11.23"),
    answer("10.78", correct = T),
    answer("10.91"),
    allow_retry = T,
    correct = "Bom trabalho!",
    incorrect = "Resposta errada"
  )
)
```

### Exercício 2

Determine a área de um círculo com diâmetro de `20` cm e atribua o resultado a um objeto chamado area_circulo. Eis a fórmula para calcular a área do círculo:

$$ A = \pi r^2 $$

```{r area-circulo, exercise=TRUE}

```

::: {#area-circulo-hint}
**Dica:** Você pode usar `pi` para que o R retorne o valor de $\pi$.
:::

```{r area-circulo-quiz}
quiz(
  question("Qual o valor encontrado para área do círculo",
    answer("1256.637"),
    answer("314.1593", correct = T),
    answer("3947.842"),
    answer("214.1593"),
    allow_retry = T,
    correct = "Parabéns!",
    incorrect = "Resposta errada"
  )
)
```

### Exercício 3

Calcule a raiz cúbica de `14` x `0.51`. Arredonde o resultado para 2 casas decimais.

```{r raiz-cubica, exercise=TRUE}

```

```{r raiz-cubica-hint}

print(3^3)

print(27^(1/3)) 
```

```{r raiz-cubica-quiz}
quiz(
  question("Qual o valor encontrado?",
    answer("1.86"),
    answer("11.19"),
    answer("363.99"),
    answer("1.93", correct = T),
    allow_retry = T,
    correct = "Muito bom!",
    incorrect = "Resposta errada"
  )
)

```

### Exercício 4

Para encontrar as soluções para uma equação do formato $ax^{2} + bx + c$, use a equação quadrática:

$x = \frac{-b \pm \sqrt{(b^2-4ac)}}{2a}$.

Ache as duas soluções para $2x^2-x-4=0$. Arredonde o resultado para duas casas decimais

```{r equacao, exercise=TRUE}

```

```{r equacao-solution}
# Definindo os coeficientes da equação
a <- 2

b <- -1

c <- -4

# Calculando as soluções da equação
x1 <- (-b + sqrt(b^2 - 4*a*c)) / (2*a)

x2 <- (-b - sqrt(b^2 - 4*a*c)) / (2*a)

## Obtendo os valores arredondados

round(x1,2)

round(x2,2)
```

```{r equacao-quiz}
quiz(
  question("Qual os valores encontrados como solução para a equação?",
    answer("1.86 e 1.69"),
    answer("1.19 e -1.69"),
    answer("1.69 e -1.19", correct = T),
    answer("1.93 e 1.18"),
    allow_retry = T,
    correct = "Ótimo trabalho!",
    incorrect = "Resposta errada"
  )
)
```

### Exercício 5

Crie um vetor chamado `peso` contendo o peso (em kg) de 4 pessoas: `80.2, 56.3, 70.5 e 60.3`. Agora, calcule a média e o desvio padrão do vetor `peso`. Arredonde os resultados para uma casa decimal.

```{r vetor-peso, exercise=TRUE}

```

```{r vetor-peso-hint}
c()

mean()

sd()

round()
```

```{r vetor-peso-quiz}
quiz(
  question("Qual os valores encontrados para a média e o desvio padrão do objeto `peso`, respectivamente?",
    answer("61.5 e 10.7"),
    answer("4.6 e 62.5"),
    answer("73.1 e 5.3"),
    answer("66.8 e 10.7", correct = T),
    allow_retry = T,
    correct = "Muito bom!",
    incorrect = "Resposta errada"
  )
)


```

### Exercício 6

Crie o objeto `altura` com os valores (em cm) para a altura das mesmas quatro pessoas que tiveram o peso registrado anteriormente: `175, 160, 165 e 172`. Agora, transforme o objeto `altura` de modo que os valores sejam apresentados em metros. Salve o resultado para um novo objeto chamado `altura_m`.

```{r vetor-altura, exercise=TRUE}

```

```{r vetor-altura-quiz}

quiz(
  question("Qual os valores encontrados para a mediana, o mínimo e o máximo do objeto `altura_m`, respectivamente?",
    answer("1.685, 1.6 e  1.75", correct = T),
    answer("168.5, 160 e 175"),
    answer("1.798, 1.6 e 1,75"),
    allow_retry = T,
    correct = "Muito bom!",
    incorrect = "Resposta errada"
  )
)

```

### Exercício 7

Sabendo que o Índice de Massa Corpórea (IMC) é obtido dividindo o peso (em kg) pela altura (em metros) ao quadrado, recrie os objetos `peso` e `altura_m` e crie um objeto `imc` com os valores do IMC de cada caso.

```{r vetor-imc, exercise=TRUE}

```

```{r veto-imc-quiz}
quiz(
  question("Qual o imc médio dessa amostra de quatro pessoas? Arredonde o resultado para três casas decimais",
    answer("23.944"),
    answer("0.002"),
    answer("23.614", correct = T),
    allow_retry = T,
    correct = "Muito bom!",
    incorrect = "Resposta errada"
  )
)


```

### Exercício 8

No R, é possível usar funções de dentro de funções. Veja o exemplo a seguir:

```{r echo=TRUE}
a <- c(2.567,3.786,4.89,5.765,9.27) 

round(mean(a), 2) 

```

Nesse código, criamos um objeto `a` e, em seguida, usamos a função `mean()` dentro da função `round()` para obter um resultado com duas casas decimais.

Com as funções `paste()` e `rep()`, em uma única linha de código, crie um objeto `questoes` com os seguintes elementos `q1,q2,q3,q4 e q5`. Em uma segunda linha, imprima o objeto `questoes` com a função `print()` para conferir o resultado.

```{r paste-rep, exercise=TRUE}

```

```{r paste-rep-solution, exercise.reveal_solution = FALSE}

questoes <- paste(rep("q",5),1:5, sep = "")
print(questoes)

```

```{r paste-rep-code-check}
grade_code()
```

::: {#paste-rep-hint}
**Dica:** Você pode encontrar orientações sobre como usar a função `paste()` [aqui](https://www.statology.org/r-paste/ "paste()") e para a fução `rep()` [aqui](https://www.statology.org/rep-function-in-r/ "rep()"). Na função `paste()`, não esqueça de definir o argumento `sep`.
:::

### Exercício 9

Agora, crie um objeto `questoes_` com os seguintes elementos: `q_1,q_2,q_3,q_4 e q_5`. Em seguida, em uma nova linha, imprima o objeto `questoes_` com a função `print()` para conferir o resultado.

```{r paste-rep_2, exercise=TRUE}

```

```{r paste-rep_2-solution, exercise.reveal_solution = FALSE}

questoes_ <- paste(rep("q",5),1:5,sep = "_")
print(questoes_)

```

```{r paste-rep_2-code-check}
grade_code()
```

### Exercício 10

Agora, tente criar novamente o objeto `questoes` com os elementos `q1,q2,q3,q4 e q5`, mas utilizando o operador pipe (`|>` ou `%>%`). Em uma segunda linha, imprima o objeto `questoes` com a função `print()`

```{r paste-rep-3, exercise=TRUE}

```

```{r paste-rep-3-solution, exercise.reveal_solution = FALSE}
questoes <- rep("q",5) |> 
  paste(1:5, sep = "")

print(questoes)
```

```{r paste-rep-3-code-check}
grade_code()
```

### Exercício 11

Crie um objeto `histograma` com 1000 valores que tenham uma distribuição normal com média 20 e desvio padrão 5. Em seguida, crie um histograma usando a função básica do R.

```{r hist, exercise=TRUE}

```

```{r hist-hint}

rnorm()

hist()

```

```{r hist-solution, exercise.reveal_solution = FALSE}
histograma <- rnorm(1000,20,5)

hist(histograma)
```

```{r hist-code-check}

grade_code()

```

### Exercício 12

Use a função `replace()` para substituir os valores do objeto `a` que estão como caracteres pelo valor numérico correspondente.

```{r replace, exercise=TRUE}
a <- c(1,2,"três",4,"cinco")

```

::: {#replace-hint}
**Dica:** Você pode encontrar orientações sobre como usar a função `replace()` [aqui](https://www.statology.org/replace-function-in-r/ "replace()")
:::

```{r replace-solution, exercise.reveal_solution = FALSE}

a <- c(1,2,"três",4,"cinco")

replace(a,c(3,5),c(3,5))
```

```{r replace-code-check}
grade_code()
```

### Exercícios 13

Extraía o nome das variáveis do banco de dados `mtcars` e salve num objeto chamado `colunas`. Em seguinda, imprima o objeto criado.

`mtcars` é um banco de dados que vem com R e apresenta dados extraídos da revista Motor Trend US de 1974 e compreende o consumo de combustível e 10 aspectos do design e desempenho de automóveis para 32 automóveis (modelos de 1973 a 1974).

Para saber mais sobre este banco use `?mtcars`. Se quiser ver outros banco de dados disponíveis no R, use a função `data()`

```{r colnames, exercise=TRUE}

```

```{r colnames-hint}

colnames()
```

```{r colnames-solution, exercise.reveal_solution = FALSE}

colunas <- colnames(mtcars)

print(colunas)
```

```{r colnames-code-check}
grade_code()
```

## Gráficos com `ggplot2`

### Introdução

Neste bloco de exercícios você vai utilizar o pacote `ggplot2` para construir os seguintes gráficos

-   Gráfico de barras
-   Histogramas

O banco de dados `starwars` vai ser utilizado para este conjunto de atividades. `starwars` contém características de 87 personagens da franquia criada por George Lucas. Para saber mais sobre esses dados, utilize:

```{r eval=FALSE, echo=TRUE}
?starwars
```

Para dar uma olhada nas variávies do bando de dados, você pode utilizar o comando a seguir

```{r eval=FALSE, echo=TRUE}
starwars |> 
  glimpse()
```

Tente executar a função `glimpse()`

```{r glimpse, exercise=TRUE}

```

### Exercício 1

Reproduza o seguinte gráfico a partir do banco de dados `starwars`:

```{r}
starwars |>
  drop_na(sex) |> 
  ggplot(aes(fct_infreq(sex), fill = sex)) +
  geom_bar() +
  theme(
    legend.position = "none"
    ) 
```

```{r bar, exercise=TRUE}


```

```{r bar-hint}

drop_na()

fct_infreq()

theme()

```

### Exercício 2

Modifique o gráfico da questão anterior, de modo que ele fique como na imagem a seguir:

```{r}
starwars |>
  drop_na(sex) |> 
  ggplot(aes(fct_infreq(sex), fill = sex)) +
  geom_bar() +
  labs(
    x = "Sexo",
    y = "Contagem"
    ) +
  theme(
    legend.position = "none"
    ) 
```

```{r bar-2, exercise=TRUE}


```

```{r bar-2-hint}
labs()
```

### Exercício 3

Reproduza o gráfico a seguir utilizando o banco de dados `starwars`:

```{r}
starwars |> 
  drop_na(sex,gender) |> 
  ggplot(aes(gender, fill = sex)) +
  geom_bar(
    position = "fill"
    ) +
  labs(
    x = "Gênero",
    y = "Contagem",
    fill = "Sexo"
  )
```













